```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz_04.ipynb")
```

# Quiz 4

## Instructions

1. After you answer a question, you should run its public tests.
1. After you answer every question, you should:
    1. Restart your kernel and clear all output
    1. Run up to the last cell
    1. Save your notebook
    1. Run the last cell to create a .zip file for Gradescope
    1. Upload this .zip file to Gradescope
    1. ***Make sure your Gradescope autograder results match your local autograder results***
1. This quiz has public and hidden tests:
    1. Public tests check your answers for correct types and shapes but may not completely check your answers
    1. Hidden tests completely check your answers but will not be available until after the due date
    1. *If all tests were public, you could quickly reverse-engineer incorrect answers that trick the autograder*
1. You may ask technical questions on Canvas Discussions, but this quiz is an individual effort.

## Packages and Settings

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

## Data

I used the following code to download and save the raw data you need for this quiz.
These are Yahoo! Finance daily data from 2011 through 2022 (inclusive).
***You should review the following code but not run it.***

```python
# import yfinance as yf
# import pandas_datareader as pdr

# (
#     yf.Tickers(tickers=['AAPL', 'TSLA'])
#     .history(period='max', auto_adjust=False, progress=False)
#     .assign(Date = lambda x: x.index.tz_localize(None))
#     .set_index('Date')
#     .loc['2011':'2022']
#     .to_csv('yahoo.csv')
# )
```

I used the following code to load the raw data you need for this quiz.
***Run the following code as-is (i.e., do not change it).***

```python
yahoo = (
    pd.read_csv(
        filepath_or_buffer='yahoo.csv',
        header=[0, 1],
        index_col=0, 
        parse_dates=True
    )
    .rename_axis(columns=['Variable', 'Ticker'])
)

yahoo.head()
```

Output:
```
Variable    Adj Close                Close           Dividends       \
Ticker           AAPL      TSLA       AAPL      TSLA      AAPL TSLA   
Date                                                                  
2011-01-03  10.017971  1.774667  11.770357  1.774667       0.0  0.0   
2011-01-04  10.070251  1.778000  11.831786  1.778000       0.0  0.0   
2011-01-05  10.152624  1.788667  11.928571  1.788667       0.0  0.0   
2011-01-06  10.144423  1.858667  11.918929  1.858667       0.0  0.0   
2011-01-07  10.217068  1.882667  12.004286  1.882667       0.0  0.0   

Variable         High                  Low                 Open            \
Ticker           AAPL      TSLA       AAPL      TSLA       AAPL      TSLA   
Date                                                                        
2011-01-03  11.795000  1.800000  11.601429  1.726667  11.630000  1.789333   
2011-01-04  11.875000  1.796667  11.719643  1.734667  11.872857  1.777333   
2011-01-05  11.940714  1.793333  11.767857  1.
```

## Questions

### Question 1

Calculate the decimal daily returns for AAPL and TSLA.
Add these return to data frame `yahoo` with outer index level name `Returns`.

_Points:_ 20

```python
_ = pd.MultiIndex.from_product([['Returns'], yahoo['Adj Close']])
yahoo[_] = yahoo['Adj Close'].pct_change()
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: yahoo has a column multi-index (+10 pts)
q1 - 2 message: yahoo has the expected column sums (+10 pts)
```

### Question 2

Calculate the monthly mean and standard deviation of returns for AAPL and TSLA from 2015 through 2022 (inclusive).
Assign this new data frame to `yahoo_m`.
Leave your means and standard deviations as daily decimal return values (i.e., do not annualize or convert to percent).

***Note: This question has 2 hidden tests worth 20 points.***

_Points:_ 40

```python
yahoo_m = (
    yahoo.loc["2015":"2022"]["Returns"]
    .groupby(pd.Grouper(freq='M'))
    .agg(['mean','std'])
)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: yahoo_m has the shape (96, 4) (+10 pts)
q2 - 2 message: yahoo_m column multi-index has the correct levels (+10 pts)
```

### Question 3

<!-- BEGIN QUESTION -->

Create portfolio of AAPL and TSLA where $w_{AAPL} = \frac{1}{3}$ and $w_{TSLA} = \frac{2}{3}$.
Assign the daily decimal returns for this portfolio to series `yahoo_p`.
Calculate the total return on this portfolio from 2014 to 2022 (inclusive) and assign this value to `totret`.
Leave `totret` as a decimal return (i.e., do not convert it to percent).

***Note: This question has 2 hidden tests worth 20 points.***

_Points:_ 40

```python
yahoo_p= yahoo["Returns"].dot([1/3,2/3])
totret = (yahoo_p["2014":"2022"].add(1).prod())-1
totret
```

Output:
```
12.408753257389012
```

```python
grader.check("q3")
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 20)
```
>>> type(yahoo.columns) is pd.core.indexes.multi.MultiIndex
True
```
failure_message: yahoo should have a column multi-index (+0 pts)
success_message: yahoo has a column multi-index (+10 pts)
```
>>> np.allclose(yahoo['Returns'].sum().values, np.array([3.06190606, 6.13513954]))
True
```
failure_message: yahoo does not have the expected column sums (+0 pts)
success_message: yahoo has the expected column sums (+10 pts)


Tests for q2 (points: 40)
```
>>> yahoo_m.shape == (96, 4)
True
```
failure_message: yahoo_m should have the shape (96, 4) (+0 pts)
success_message: yahoo_m has the shape (96, 4) (+10 pts)
```
>>> (yahoo_m.columns.get_level_values(0).to_list() == ['AAPL', 'AAPL', 'TSLA', 'TSLA']) & (yahoo_m.columns.get_level_values(1).to_list() == ['mean', 'std', 'mean', 'std'])
True
```
failure_message: yahoo_m column multi-index has the incorrect levels (+0 pts)
success_message: yahoo_m column multi-index has the correct levels (+10 pts)


Tests for q3 (points: 40)
```
>>> (yahoo_p.shape == (3020,)) & ((type(totret) is float) | (type(totret) is np.float64))
True
```
failure_message: yahoo_p should have the shape (3020,) and totret should be a float (+0 pts)
success_message: yahoo_p has the shape (3020,) and totret is a float (+10 pts)
```
>>> yahoo['Returns'].sum().min() < yahoo_p.sum() < yahoo['Returns'].sum().max()
True
```
failure_message: sum of yahoo_p should be between sums of AAPL and TSLA (+0 pts)
success_message: sum of yahoo_p is between sums of AAPL and TSLA (+10 pts)
