```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz_05.ipynb")
```

# Quiz 5

## Instructions

1. After you answer a question, you should run its public tests.
1. After you answer every question, you should:
    1. Restart your kernel and clear all output
    1. Run up to the last cell
    1. Save your notebook
    1. Run the last cell to create a .zip file for Gradescope
    1. Upload this .zip file to Gradescope
    1. ***Make sure your Gradescope autograder results match your local autograder results***
1. This quiz has public and hidden tests:
    1. Public tests check your answers for correct types and shapes but may not completely check your answers
    1. Hidden tests completely check your answers but will not be available until after the due date
    1. *If all tests were public, you could quickly reverse-engineer incorrect answers that trick the autograder*
1. You may ask technical questions on Canvas Discussions, but this quiz is an individual effort.

## Packages and Settings

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

## Data

I used the following code to download and save the raw data you need for this quiz.
***You should review the following code but not run it.***

```python
# import yfinance as yf

# (
#     yf.Tickers(tickers=['AAPL', 'TSLA'])
#     .history(period='max', auto_adjust=False, progress=False)
#     .assign(Date = lambda x: x.index.tz_localize(None))
#     .set_index('Date')
#     .rename_axis(columns=['Variable', 'Ticker'])
#     .to_csv('yahoo.csv')
# )
```

## Questions

### Question 1

Read the daily returns data for AAPL and TSLA saved in file `yahoo.csv`
Calculate the decimal daily returns for AAPL and TSLA from July 2010 through February 2023 (inclusive) and assign to data frame `returns`.

_Points:_ 28

```python
returns = (
    pd.read_csv('yahoo.csv', header=[0,1], index_col=0)["Adj Close"]
    .pct_change()
    .loc["2010-07":"2023-03"]
)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: returns has the right shape (+14 pts)
q1 - 2 message: returns has the right column sums (+14 pts)
```

### Question 2

Calculate the 252-trading-day rolling correlations between AAPL and TSLA daily returns.
Use all available returns in data frame `returns`, and require 252 observations per correlation calculation. 
Leave your rolling correlations as decimals (i.e., do not convert to percent), and assign them to series `rhos`.

***Note: This question has 3 hidden tests worth 18 points.***

_Points:_ 24

```python
rhos = returns["AAPL"].rolling(252).corr(returns["TSLA"])
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: rhos has the right shape (+6 pts)
```

### Question 3

<!-- BEGIN QUESTION -->

Aggregate the daily returns in data frame `returns` to monthly total returns.
Save these monthly total returns to data frame `returns_m`.

Calculate the 12-month rolling correlation between AAPL and TSLA monthly returns.
Use all available returns in data frame `returns_m`, and require 12 observations per correlation calculation.
Leave your rolling correlations as decimals (i.e., do not convert to percent), and assign them to series `rhos_m`.

***Note: This question has 6 hidden tests worth 36 points.***

_Points:_ 48

```python
returns.index = pd.to_datetime(returns.index)
returns_m = returns.add(1).resample("M").prod().sub(1)
rhos_m = returns_m["AAPL"].rolling(12).corr(returns_m["TSLA"])
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 1 message: returns_m has the right shape (+6 pts)
q3 - 2 message: rhos_m has the right shape (+6 pts)
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: returns has the right shape (+14 pts)
    q1 - 2 message: returns has the right column sums (+14 pts)

    q2 results: All test cases passed!
    q2 - 1 message: rhos has the right shape (+6 pts)

    q3 results:
        q3 - 1 message: returns_m has the right shape (+6 pts)

        q3 - 1 result:
            ✅ Test case passed

        q3 - 2 message: rhos_m has the wrong shape (+0 pts)

        q3 - 2 result:
            ❌ Test case failed
            Trying:
                rhos_m.shape == (152,)
            Expecting:
                True
            **********************************************************************
            Line 1, in q3 1
            Failed example:
                rhos_m.shape == (152,)
            Exception raised:
                Traceback (most recent ca
```

Tests for q1 (points: 28)
```
>>> returns.shape == (3187, 2)
True
```
failure_message: returns has the wrong shape (+0 pts)
success_message: returns has the right shape (+14 pts)
```
>>> np.allclose(returns.sum().values, np.array([3.47652148, 6.94111853]))
True
```
failure_message: returns has the wrong column sums (+0 pts)
success_message: returns has the right column sums (+14 pts)


Tests for q2 (points: 24)
```
>>> rhos.shape == (3187,)
True
```
failure_message: rhos has the wrong shape (+0 pts)
success_message: rhos has the right shape (+6 pts)


Tests for q3 (points: 48)
```
>>> returns_m.shape == (152, 2)
True
```
failure_message: returns_m has the wrong shape (+0 pts)
success_message: returns_m has the right shape (+6 pts)
```
>>> rhos_m.shape == (152,)
True
```
failure_message: rhos_m has the wrong shape (+0 pts)
success_message: rhos_m has the right shape (+6 pts)
