```python
# Initialize Otter
import otter
grader = otter.Notebook("quiz_06.ipynb")
```

# Quiz 6

## Instructions

1. ***Do not change any file names or relative locations***
1. After you answer a question, you should run its public tests.
1. After you answer every question, you should:
    1. Restart your kernel and clear all output
    1. Run up to the last cell
    1. Save your notebook
    1. Run the last cell to create a .zip file for Gradescope
    1. Upload this .zip file to Gradescope
    1. ***Make sure your Gradescope autograder results match your local autograder results***
1. This quiz has public and hidden tests:
    1. Public tests check your answers for correct types and shapes but may not completely check your answers
    1. Hidden tests completely check your answers but will not be available until after the due date
    1. *If all tests were public, you could quickly reverse-engineer incorrect answers that trick the autograder*
1. You may ask technical questions on Canvas Discussions, but this quiz is an individual effort.

## Packages and Settings

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
import scipy.optimize as sco
```

## Data

I used the following code to download and save the data you need for this quiz.
***You should review the following code but not run it.***

```python
# import yfinance as yf

# (
#     yf.Tickers(tickers=['TSLA', 'BAC', 'FRC', 'F', 'AMD'])
#     .history(period='max', auto_adjust=False, progress=False)
#     .assign(Date = lambda x: x.index.tz_localize(None))
#     .set_index('Date')
#     .rename_axis(columns=['Variable', 'Ticker'])
#     .to_csv('yahoo.csv')
# )
```

## Questions

### Question 1

Read the daily price data for AMD, BAC, F, FRC, and TSLA saved in file `yahoo.csv` and assign to data frame `yahoo`.
Calculate the daily decimal returns for AMD, BAC, F, FRC, and TSLA from March 2020 through February 2023 (inclusive) and assign to data frame `returns`.

_Points:_ 40

```python
returns = pd.read_csv("yahoo.csv", header= [0,1] , index_col= [0])["Adj Close"].pct_change().loc["2020-03":"2023-03"]

```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: returns has the right shape (+15 pts)
q1 - 2 message: returns has the right column sums (+15 pts)
q1 - 3 message: returns has the right column names and order (+10 pts)
```

### Question 2

Use the daily decimal returns in data frame `returns` to calculate the minimum-variance portfolio weights.
Assume long-only portfolio weights, so all portfolio weights are non-negative.
Assign these minimum-variance portfolio weights to NumPy array `weights_2`.
***Note: This question has 4 hidden tests worth 24 points.***

_Points:_ 36

```python
def port_vol(x, r, ppy=252):
    return np.sqrt(ppy) * r.dot(x).std()
```

```python
a = sco.minimize(
    fun=port_vol, 
    x0=np.ones(returns.shape[1]) / returns.shape[1], 
    args=(returns,), 
    bounds=[(0,1) for _ in returns], 
    constraints=(
        {'type': 'eq', 'fun': lambda x: x.sum() - 1} 
    )
)
weights_2 = a["x"]
print(weights_2)
```

Output:
```
[0.13669878 0.33293334 0.08658019 0.39122437 0.05256332]
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 1 message: weights_2 has the right shape (+6 pts)
q2 - 2 message: weights_2 values fall between 0 and 1 (+6 pts)
```

### Question 3

Calculate the minimum-variance portfolio weights as above in question 2, but allow long and short portfolio weights.
Constrain the sum of the long portfolio weights to be less or equal to 130% or 1.3.
Assign these minimum-variance portfolio weights to NumPy array `weights_3`.
***Note: This question has 4 hidden tests worth 16 points.***

_Points:_ 24

```python
b = sco.minimize(
    fun=port_vol, 
    x0=np.ones(returns.shape[1]) / returns.shape[1],
    args=(returns,), 
    bounds = [(-0.7,1.3)for _ in returns],
    constraints=(
         {'type': 'ineq', 'fun': lambda x:x[x>0].sum()-1.3 }
    )
)
weights_3 = b["x"]
weights_3[weights_3>0]
```

Output:
```
array([0.17763965, 0.43282895, 0.11268624, 0.50832488, 0.06852028])
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 1 message: weights_3 has the right shape (+4 pts)
q3 - 2 message: weights_3 values fall between -1 and 2 (+4 pts)
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q1 results: All test cases passed!
    q1 - 1 message: returns has the right shape (+15 pts)
    q1 - 2 message: returns has the right column sums (+15 pts)
    q1 - 3 message: returns has the right column names and order (+10 pts)

    q2 results: All test cases passed!
    q2 - 1 message: weights_2 has the right shape (+6 pts)
    q2 - 2 message: weights_2 values fall between 0 and 1 (+6 pts)

    q3 results:
        q3 - 1 message: weights_3 has the wrong shape (+0 pts)

        q3 - 1 result:
            ❌ Test case failed
            Trying:
                weights_3.shape == (5,)
            Expecting:
                True
            **********************************************************************
            Line 1, in q3 0
            Failed example:
                weights_3.shape == (5,)
            Exception raised:
        
```

Tests for q1 (points: None)
```
>>> returns.shape == (755, 5)
True
```
failure_message: returns has the wrong shape (+0 pts)
success_message: returns has the right shape (+15 pts)
```
>>> np.allclose(returns.sum().values, np.array([1.013422  , 0.50552809, 1.01683573, 0.4575595 , 2.27742942]))
True
```
failure_message: returns has the wrong column sums (+0 pts)
success_message: returns has the right column sums (+15 pts)
```
>>> (returns.columns == np.array(['AMD', 'BAC', 'F', 'FRC', 'TSLA'])).all()
True
```
failure_message: returns has the wrong column names or order (+0 pts)
success_message: returns has the right column names and order (+10 pts)


Tests for q2 (points: None)
```
>>> weights_2.shape == (5,)
True
```
failure_message: weights_2 has the wrong shape (+0 pts)
success_message: weights_2 has the right shape (+6 pts)
```
>>> ((0 <= weights_2) & (weights_2 <= 1)).all()
True
```
failure_message: weights_2 values do not fall between 0 and 1 (+0 pts)
success_message: weights_2 values fall between 0 and 1 (+6 pts)


Tests for q3 (points: None)
```
>>> weights_3.shape == (5,)
True
```
failure_message: weights_3 has the wrong shape (+0 pts)
success_message: weights_3 has the right shape (+4 pts)
```
>>> ((-1 <= weights_3) & (weights_3 <= 2)).all()
True
```
failure_message: weights_3 values do not fall between -1 and 2 (+0 pts)
success_message: weights_3 values fall between -1 and 2 (+4 pts)
